---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## what is docker

```yaml
type: NormalExercise
key: e8c1edbe67
lang: bash
xp: 100
skills: 2
```

Docker is the most usable container technology, but what is a container ? in the `linux` world the container feature is not new, it's introduced in the linux kernel as `LXC : LinuX Container`.

`@instructions`
* Check if Docker is running with `systemctl status docker`
* get the Docker version `docker version`
* get the help commands with `docker --help`

`@hint`
soon

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{shell}

```

`@solution`
```{shell}

```

`@sct`
```{python}

```
