---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## what is docker

```yaml
type: NormalExercise
key: e8c1edbe67
lang: shell
xp: 100
skills: 2
```

Docker is the most usable container technology, but what is a container ? in the `linux` world the container feature is not new, it's introduced in the linux kernel as `LXC : LinuX Container`.

`@instructions`


`@hint`


`@pre_exercise_code`
```{shell}

```

`@sample_code`
```{shell}

```

`@solution`
```{shell}

```

`@sct`
```{shell}

```
